# COGS 108 - Assignment 1: Git & `pandas`

This assignment covers git usage, the basics of data wrangling, and working with data in `pandas`. 

This assignment is out of 6 points, worth 6% of your grade.

**PLEASE DO NOT CHANGE THE NAME OF THIS FILE.**

**PLEASE DO NOT COPY & PASTE OR DELETE CELLS INCLUDED IN THE ASSIGNMENT.** (Note that you *can* add additional cells, if you want to test things out.)

## How to complete assignments

Whenever you see:

```python
# YOUR CODE HERE
raise NotImplementedError()
```

You need to **replace (meaning: delete) these lines of code with some code that answers the questions** and meets the specified criteria. Make sure you remove the 'raise' line when you do this (or your notebook will raise an error, regardless of any other code, and thus fail the grading tests).

You should write the answer to the questions in those cells (the ones with `# YOUR CODE HERE`), but you can also add extra cells to explore / investigate things if you need / want to. 

### How to Score Points

Any cell with `assert` statements in it is a test cell. You should not try to change or delete these cells. Note that there might be more than one assert that tests a particular question, and there may be some additional hidden `assert` statements to test your code.

Reading assert errors if and when they fail is a great way to debug your code and find out what went wrong.

Note that some cells, including the test cells, may be read only, which means they won't let you edit them. If you cannot edit a cell - that is normal, and you shouldn't need to edit that cell.

All outside packages/modules that will be used will be specified. You may not use other libraries in the assignments.

Finally, note that questions have points as specified in the detailed instructions. 

## Part 1: Git and Github

In D1, we learned how to make commits, create branches, and open pull requests. Now, we'll take the next step by learning how to handle situations that come up when working with others on shared code. 

Merge conflicts arise when two people make changes to the same part of a file -- Git needs your help to decide which version to keep. You will also learn how to review pull requests, an important part of effective collaboration on GitHub projects.

In this section, we will go over common GitHub tasks. Complete the following tutorials: 
- [Review Pull Requests](https://github.com/skills/review-pull-requests)
- [Resolve Merge Conflicts](https://github.com/skills/resolve-merge-conflicts)

Open the links above and follow the instructions in the README repository.

The GitHub Skills tutorials use a cool trick to walk you through the steps... every time you complete a task, it uses a GitHub Actions automation to trigger an update of the instructions in the README to show you the next thing you need to complete.

**Some Important Notes** 

- Follow every step exactly as written in the GitHub Skills tutorial README.
- **When starting the tutorial,click the green 'Start Course' button you can see in the README**. Do not use any other way to start the tutorial.  You will lose points if you do; and I'm not listing the other ways because somebody would accidentally think I was saying to use those methods.
- When you press Start Course you will get a copy of the repo in your own GitHub repositories.  **KEEP THE DEFAULT REPO NAME!!**  The Skills repos we are doing will be by default called `skills-review-pull-requests` and `skills-resolve-merge-conflicts` in your GitHub.  If you rename your repo something else, you will lose points from the autograder.
- What if you already did these skills tutorials last week/quarter/whatever?  If you already have a repo called `skills-review-pull-requests` in your GitHub you won't be allowed to create a new one with the same name.  If you feel comfortable with your skills just make sure the old copy of the tutorial actually was finished that time.  If you want to refresh your skills, delete your old copy of the repo from your GitHub and restart the tutorial.
- The instructions in the README suggest you should open two windows on the your GitHub tutorial repo... one to keep the instructions visible to you and another to actually execute the steps.  This is very important... once you try to do something (commit, pull request, etc) you will be on a new webpage and the instructions disappear.  Its hard to keep all the steps in your head, so its crticial to have a 2nd window open with the instructions for when you need to double check things

### Q1 (1 point)

Once you've completed the exercise, please provide us your Github username to validate the completion of the Github Skills Introduction. In the cell below, provide your Github username (not the email address you used to make your account) as a string in the variable github_username.

<!-- grade_id: cell-4559f70354cbbe79, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e4df724c7ced46e6, points: 0, grade: true, solution: false -->
```python
# Autograder tests
# If this cell executes with no output, you passed the tests!
assert github_username, 'Github username is not entered'
assert type(github_username) == str, 'Username is not a string variable'

import requests
url = 'https://github.com/' + github_username
response = requests.get(url)
assert response.status_code == 200
```

<!-- grade_id: cell-75e562ea892caaaf, points: 0.5, grade: true, solution: false -->
```python
# Do not edit this cell. This cell is for testing purposes
```

<!-- grade_id: cell-881e39752ea77110, points: 0.5, grade: true, solution: false -->
```python
# Do not edit this cell. This cell is for testing purposes
```

## Part 2: Merging Conflicts on DataHub (0.5 points)

In this part of the assignment, you’ll get hands-on practice resolving merge conflicts in a Jupyter notebook.  


You’ve already learned about merge conflicts conceptually and worked with them using GitHub in the previous Git skills activity. Now, we’ll extend that knowledge into the environment you’ll be using most often for your projects: Datahub.  This exercise will use the the Git GUI built in to Datahub

Why this matters: merge conflicts are an unavoidable part of collaborative coding. Whether you’re working on a shared notebook, script, or dataset, it’s essential to understand how to identify, manage, and resolve conflicts safely. Merge conflicts are also more difficult to deal with in Jupyter notebooks than in regula text files.  This exercise will help you gain the confidence to do this in your project repos.

Follow the instructions on this [repository](https://github.com/COGS108/GitExercise_Conflicts) to get started.

BE ABSOLUTELY SURE you do ALL of the git operations in the tutorial  as instructed. If you don't make some changes requested, or choose a different set of changes you may lose points on the autograder


<!-- grade_id: cell-9a43850d2bb8bb55, points: 0.5, grade: true, solution: false -->
```python
# Do not edit this cell. This cell is for testing purposes
```

## Part 3: Pandas 

This part of the assignment is focused on some practice with data set manipulation with `pandas`. 

You will absolutely use `pandas` throughout this course, and, if you continue to use Python for working with data in the future, will continue to use this package. This means that while you could Google and look up every `pandas` function/method when doing assignments, it's best to memorize the basics. So, do your best to understand each operation you carry out here, as you'll almost certainly see it again. And, try to commit the most commonly-used functions/methods (those being presented here) to memory. It will save future you a ton of time. The more comfortable you are with `pandas`, the easier the rest of the assignments and your final project will be for you!

Finally, there are a ton of great `pandas` resources and documentation out there, including [10 minutes to `pandas`](https://pandas.pydata.org/pandas-docs/stable/user_guide/10min.html). Feel free to use them! However, many of the examples in this assignment were inspired by `pandas` tips and tricks: https://www.dataschool.io/python-pandas-tips-and-tricks/. Feel free to check out all the additional tips and tricks there, many of which are *not* covered here. 

### Q2: Import (0.25 points)

Import the following packages using their common shortened name found in parentheses:

- numpy (np)
- pandas (pd)

<!-- grade_id: cell-2792d95352c9d6d2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-826f554fa0ea7fb9, points: 0.25, grade: true, solution: false -->
```python
# Do not edit this cell. This cell is for testing purposes
```

**Run the following cell code to make things throughout the rest of this workbook a little prettier.** (Note: You don't have to edit code here, but are free to and see what changes to be sure you understand each line.)

```python
# Configure libraries

# Don't display too many rows/cols of DataFrames
pd.options.display.max_rows = 10
pd.options.display.max_columns = 10

# Round decimals when displaying DataFrames
pd.set_option('display.precision', 2)
```

### Q3: Read data into `pandas` (0.25 points)

The `pandas` package has a number of helpful import functions (read_csv(), read_json(), read_html(), etc.)

For this assignment we'll use a CSV file from [FiveThirtyEight](https://fivethirtyeight.com). The summarized data in the file we'll use here were used in the article [*The Ultimate Halloween Candy Power Ranking*](https://fivethirtyeight.com/features/the-ultimate-halloween-candy-power-ranking/). To collect these data, candies were pitted against one another with participants having to choose which they liked better.

To use the data, you'll have to read it in as a Dataframe. Use the appropriate pandas import function to read in the file from the following repo: https://github.com/fivethirtyeight/data/tree/master/candy-power-ranking.

Note that URLs can be read in directly using `pandas`. The URL you'll want to use to read the file in is: https://raw.githubusercontent.com/fivethirtyeight/data/master/candy-power-ranking/candy-data.csv. 

Store this in the variable `df`.

<!-- grade_id: cell-72a17dfe01a43e26, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e27e59a1e849e901, points: 0.25, grade: true, solution: false -->
```python
assert isinstance(df, pd.DataFrame)
assert df.shape == (85, 13)
```

Let's get a sense of the information we have in our dataset

```python
# take a look at the information stored in df
df
```

A full description of these data can be found [here](https://github.com/fivethirtyeight/data/tree/master/candy-power-ranking). From that link, we'll include a description of each variable (column) here:

Header | Description
-------|------------
chocolate | Does it contain chocolate?
fruity | Is it fruit flavored?
caramel | Is there caramel in the candy?
peanutalmondy | Does it contain peanuts, peanut butter or almonds?
nougat | Does it contain nougat?
crispedricewafer | Does it contain crisped rice, wafers, or a cookie component?
hard | Is it a hard candy?
bar | Is it a candy bar?
pluribus | Is it one of many candies in a bag or box?
sugarpercent | The percentile of sugar it falls under within the data set.
pricepercent | The unit price percentile compared to the rest of the set.
winpercent | The overall win percentage according to 269,000 matchups.

Note that many of the variables store binary information, where 0 means False and 1 means True. For example, for a Snickers bar, `chocolate` would have the value 1 (becuase it *is* chocolate), while `fruity` for the same candy would store 0 (because Snickers bars are *not* fruity).

### Q4: Data Summary

After reading in a dataset, it's often helpful to get a quick sense of what information is stored in the dataset. Here you'll carry out a number of operations that you'll carry out on almost every dataset you read into Python.

### Q4a: `shape` (0.25 points)
One way to do that is to determine the `shape` of the dataset. This reports the number of rows and columns in the dataset. Store the shape of the candy dataset in the variable `candy_shape`.

<!-- grade_id: cell-897cb56135b5c5dd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a5fcebda6222ba02, points: 0.25, grade: true, solution: false -->
```python
# Do not edit this cell. This cell is for testing purposes
```

### Q4b: `describe`

`pandas` also has a very helfpul `describe` method which will provide some helpful summary statistics about your dataframe.

Use the `describe` method to calculate and display these summary statistics (You do *not* have to store the output):

<!-- grade_id: cell-a0e3c44a4d61378d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Q4c: `value_counts()` (0.25 points)

Using the `value_counts()` method, determine how many different possible values there are for the `chocolate` series in the `df` DataFrame and how many observations fall into each. 

Store the output in the object `chocolate_values`.

Take a look at the output. Be sure you understand whether or not there are more chocolate (`chocolate == 1`) or nonchocolate candies (`chocolate == 0`) in the dataset from the output.

<!-- grade_id: cell-c3bf78a0385d82cd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

chocolate_values
```

<!-- grade_id: cell-0689de8258624ecf, points: 0.25, grade: true, solution: false -->
```python
# Do not edit this cell. This cell is for testing purposes
```

### Missing Data

There are a number of different ways in which to determine whether or not data are missing. The most common approaches are summarized here:

```python
# Calculate % of missing values in each column:
df.isna().mean()

# Drop columns with any missing values:
df.dropna(axis='columns')

# Drop columns in which more than 10% of values are missing:
df.dropna(thresh=len(df)*0.9, axis='columns')

# Want to know the *count* of missing values in a DataFrame?
df.isna().sum().sum()

# Just want to know if there are *any* missing values?
df.isna().any().any()
df.isna().any(axis=None)
```

Run the following cell and interpret the output:

```python
# determine which columns have missing data
df.isna().any()
```

### Q4d: isna() (0.50 points)
How many variables have missing data in this dataset? Store the value in the variable var_missing:

<!-- grade_id: cell-342b2367989d2634, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-542821c0cb36ac33, points: 0.25, grade: true, solution: false -->
```python
assert isinstance(var_missing, (np.int64, int))
print(var_missing)
```

### Q5: `replace` (0.5 points)

To replace current values with different values dynamically, `replace` is very helpful.

To demonstrate this, if you wanted to replace the zeroes and ones in the current dataset with string values, you could do this using `replace`.  

In `df`, replace `0` with `"is not"` and `1` with `"is"`.   NOTE: we are replacing integers with strings! 

Store this output in `df_string`.

That is, where the current `df` is showing a 1 for 100 Grand in the chocolate column, the new `df_string` will show 100 Grand "is" chocolate in that same location.

<!-- grade_id: cell-05cdaa62a575a340, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-879057f9f448f7da, points: 0.5, grade: true, solution: false -->
```python
assert df_string.shape == df.shape

```

## Q6: `groupby` (0.25 points)

Group df by the candy’s chocolate status (0: is not, 1: is), and calculate the mean value of pricepercent for each group. Assign this to the dataframe 'df_grouped'. 

<!-- grade_id: cell-17fa4a9531ca2a88, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-aecf41787e90b8b4, points: 0.25, grade: true, solution: false -->
```python
print(df_grouped)
```

### Q6a: Cost Interpretation (0.25 points)

Looking at the output in the previous cell, which type of candy costs more, non-chocolate or chocolate candy? Store your response in `candy_cost`. For example, if you thought the answer was A, you would submit `candy_cost = 'A'`.

- A) candy containing chocolate costs more
- B) candy without chocolate costs more
- C) candy without and containing chocolate cost the same

<!-- grade_id: cell-5cf8c0eb59ef30aa, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7928c671100283b4, points: 0.25, grade: true, solution: false -->
```python
assert isinstance(candy_cost, str)
assert candy_cost in ['A', 'B', 'C']
```

### Q6b: specific column aggregation (0.25 points)

Carrying out a similar operation, group `df` by `fruity`, and then calculate the minimum, average (mean), and maximum values, **for the `sugarpercent` column only**. 

For this, fruity should be the name for the indices and your column names should be 'min', 'mean', and 'max'.

Store this result in `sugar_summary`.

#### Note: the `pandas` documentation [here](https://pandas.pydata.org/pandas-docs/version/0.22/generated/pandas.core.groupby.DataFrameGroupBy.agg.html) may help you if you're struggling to figure out the syntax

<!-- grade_id: cell-4f443afae517bca7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# look at output
sugar_summary
```

<!-- grade_id: cell-64a031270c0f5741, points: 0.25, grade: true, solution: false -->
```python
assert type(sugar_summary) == pd.DataFrame
assert sugar_summary.index.name == 'fruity'
assert 'min' in [j for i in sugar_summary.columns for j in i] and len(sugar_summary.columns) == 3 or 'min' in sugar_summary.columns
```

### Q6c: Sugar Interpretation (0.25 points)
Looking at the output in Q9b and considering what you know about the dataset thus far, which type of candy has the higher average sugar percentile, fruity candy or non-fruity candy? Store your response in candy_sugar. For example, if you thought the answer was A, you would submit candy_sugar = 'A'.

- A) non-fruity candy has the higher average sugar percentile
- B) fruty candy has higher average sugar percentile
- C) non-fruity and fruity have the same average sugar percentile

<!-- grade_id: cell-6f4882bb9714dcf0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ce659911f405a669, points: 0.25, grade: true, solution: false -->
```python
assert isinstance(candy_sugar, str)

```

## Q7: Column Operations 

Regardless of the dataset you're working with, you'll want to be *very* familiar with a few common data manipulations when wrangling data, each of which is described below:

Manipulation | Description
-------|------------
**select** | select which columns to include in dataset
**filter** | filter dataset to only include specified rows
**mutate** | add a new column based on values in other columns
**groupby** | group values to apply a function within the specified groups
**summarize** | calculate specified summary metric of a specified variable
**arrange** | sort rows ascending or descending order of a specified column
**merge** | join separate datasets into a single dataset based on a common column

We'll review how to carry out each of these in `pandas` and you'll get practice with each. First, let's load in a new dataset.

**For this part, we will be working with the same dataset from D2. Load in the CSV file at the URL:** https://raw.githubusercontent.com/fivethirtyeight/data/master/steak-survey/steak-risk-survey.csv **Assign the resulting dataframe to the variable 'data'.**

<!-- grade_id: cell-37f8a22aa5bb0382, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f959a65d4292411c, points: 0, grade: true, solution: false -->
```python
assert isinstance(data, pd.DataFrame)
assert data.shape == (551, 15)
```

Let's take a look at the data to remind ourselves what we are working with.

```python
data
```

Often we're only interested in working with some columns in a larger dataset and it's helpful to filter the dataset so that it only includes the columns we want. In these cases, we *select* the columns we're interested in by specifying these columns and returning the output.

In pandas, there are a number of ways in which this can be done, but the most common is to specify a list of columns you want by name and in the order you want.

For example, for a dataframe (`df`) with five columns ('A', 'B', 'C', 'D', 'E'), you could select columns 'C', 'B' and 'A' (in that order), using the following:

```python
df[['C', 'B', 'A']]
```

**Run the following cell to clean up our new dataframe.** Similarly to D2, we will 1) remove some undesirable rows and columns, 2) rename them to a shorter form, and 3) remove rows that contain all missing data.

```python
#drop first two columns and first row 
data = data.iloc[1:,2:]

#rename columns
data.columns = ['smoking','alcohol','gambling',
 'skydiving','speeding', 'cheated',
 'steak', 'steak_preference','gender',
 'age', 'income', 'education', 'region']

#drops NaNs
data = data.dropna(how='all')
```

## Q7a: Selecting columns and reordering columns (0.25 points)

For our risk behavior dataset, we're ultimately interested in determining the relationship between various risk-taking behaviors and demographic characteristics.

To get just the data we're interested in, return a DataFrame with the columns (in this exact order):
gambling, skydiving, speeding, cheated, age, income, education, and region.

Store this output in `data` (overwriting the previous data stored in `data`).

<!-- grade_id: cell-1af92717790016eb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4cd0a94131c943cb, points: 0.25, grade: true, solution: false -->
```python
assert data.shape == (541, 8)
```

```python
data
```

### Q7b: Adding new columns (0.25)

Often, when working with data, we need to add additional columns *that use information from existing columns*.

To do this, we typically use one of two approaches, summarized generally here.  If, e.g. we needed a new column with the name 'new_col' where each row was a constant value muliplied by the corresponding row of 'col'... then we could equivalently use either of these:

1. using the method `assign()` 
```python
df = df.assign(new_col = df['col'] * val)
```
2. using the new column name on the left-hand side of the equals sign
```python
df['new_col'] = df[col] * val
```

In our risk behavior dataset, we want to explore whether individuals engage in multiple types of risky behavior — specifically, gambling and speeding. We’ll create a new column called 'risk_combo' that **counts** the number of these two behaviors each person reports. So this could look something like

| gambling | speeding | risk_combo |
|----------|----------|------------| 
|   Yes    |  Yes     |  2         |
|   Yes    |  No      |  1         |
|   No     |  Yes     |   1        |
|   No     |  No      |  0         |

To do this, think about how you could:
- Compare the entries in the gambling and speeding columns to "Yes",
- Combine those results across each row,
- And use .assign() to store that result as a new column called 'risk_combo'.

(Hint: you may find it useful to select both columns together, check if they equal "Yes", and then use a method that adds up the True values across each row.)

<!-- grade_id: cell-2dede0f89f09017a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e5db3e8d704701b3, points: 0.25, grade: true, solution: false -->
```python
# Do not edit this cell. This cell is for testing purposes
```

### Data Types

Lets inspect the different data types present in the DataFrame 'data' to understand which columns contain text responses

<!-- grade_id: cell-b8a44349c3510826, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Everything here is of type object...  in python an object is the most generic kind of variable.  Basically everything is an object in this language. Frequently (but not always!) in tabular data it means this column is a text column of type 'str'

If, for instance, one of these variables held numbers then it would have a type like int or float or similar

**Oh no! Why would our data still contain variables like 'age' and 'income', which appears to be numeric?** That’s because, in this dataset, both age and income are stored as categorical strings — representing ranges (e.g., "18–29", "30–44", "45–60", "> 60") or as formatted text values (e.g., "24,999", "99,999") instead of true numeric data.

To perform any meaningful numerical analysis (like correlations or averages), **we’ll need to convert these string ranges into numeric values.**

## Q7c replace values (0.25 points)

We'll be working with the variable `age`. Let's replace the age ranges with numeric midpoint values using `replace()`. The numeric values will be stored in a new column called `age_numeric`. 

`replace()`can take a dictionary as an argument that maps old values to new values. `replace()` doesn't operate on the original; it creates a copy of the series or dataframe. So if you want to keep your hard work you must assign the output to a variable.

Create a dictionary named `age_map` that maps the age ranges to representative midpoint values. Use the following values in your `age_map` replace the age range strings in our dataset: 
- '18-29': 24
- '30-44': 37
- '45-60': 52
- '> 60': 65

Its important to note that we are mapping from a text variable to a numeric value.  That means each map element needs to be a string on the left side (dictionary key) and an integer on the right side (dictionary value)

Assign the outcome or your operations here to a new column called `age_numeric` in the `data` dataframe

<!-- grade_id: cell-aa97c384da457265, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4e25ba6b7b2eb549, points: 0.25, grade: true, solution: false -->
```python
assert "age_numeric" in data.columns
assert np.issubdtype(data["age_numeric"].dtype, np.number)
```

### Q7d: neither (0.25 points)

Take a look at your risk_combo column. How many individuals engage in neither gambling nor speeding?

In other words, how many rows have risk_combo == 0?

Store this value in the variable neither.

<!-- grade_id: cell-73e52457ca151b02, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0fba50f0683cd07d, points: 0.25, grade: true, solution: false -->
```python
print(neither)
assert neither >= 0 and neither <= len(data)

```

## Q8: Row Operations 

In addition to column operations, filtering rows to only include those of interest to your analysis can be very important. For example, in our risk behavior dataset, we might only want to analyze individuals who engage in at least one risky behavior (gambling or speeding).

To do this, we have to *filter* the data, using **slicing**.

### slicing:  `loc` and `iloc`

Subsets of data from a DataFrame can also be extracted through the use of slicing. There are two general approaches to slicing:

Manipulation | Description
-------|------------
`loc` | label-based indexing 
`iloc` | integer-based indexing


For example, if we had a dataframe like: 
 **Index**| **A** | **B** | **C**
----| --- | ---- | ---
**Al** | 1 | 2 | 3
**Bob** | 3 | 4 | 5
**Cat** | 6 | 7 | 8

You could use `loc`: 

```python
df.loc['Al':'Bob', ['A', 'B']]
```

This would return the first two rows and the two columns  `['A', 'B']`, ie

 **Index**| **A** | **B** 
----| --- | ---- 
**Al** | 1 | 2 
**Bob** | 3 | 4 

We can also pass in a *conditional statement* to `loc`, and this operation will return a dataframe only with the rows that satisfy this condition.

Alternatively, integer-based indexing could be used with `iloc` (where the `i` stands for index):

```python
df.iloc[0:5, 0:3]
```

This would return the first five rows and first THREE columns of the dataframe `df`. (As a reminder: when indicating ranges in Python, the final value is not included in what is returned. So, this returns the zero-th through the 4th indices. Index 5 is not included in the output)

Python also uses *zero-based indexing* which means the first element is indexed as zero, the second has index 1, and so on.

### Q8a: Slicing (0.25 points)

The concept of slicing can be combined with conditional statements to select only the rows that meet a certain criterion.

How many individuals engage in BOTH gambling AND speeding? 


So we have to identify all the dual gambling/speeding people. Let's place them in a datafram called `both`. 

Let’s figure out which region the first such individual is from. To do this we need to slice `both` to get just the first person in it and just their region.  Place that string in the variable `region_name`


<!-- grade_id: cell-3b770e4a0e2b84e1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-153ac11204d89dd8, points: 0.25, grade: true, solution: false -->
```python
assert isinstance(region_name, str)
assert type(both) == pd.DataFrame
assert both.shape[0] == 237

```

### Q8b: Filtering (0.25 points)

Using the concept of slicing with conditionals, filter the DataFrame `data` to only include rows where the individual engages in at least one risky behavior:

This includes anyone who gambles, speeds, OR does both.

Store this filtered DataFrame back in `data`.

<!-- grade_id: cell-f103daaf49474dc4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ea59146dbcf1d259, points: 0.25, grade: true, solution: false -->
```python
assert data.shape[0] == 500
assert ((data['gambling'] == "Yes") | (data['speeding'] == "Yes")).all()
```

## Q9: arrange rows

While you don't always care about what order the rows are sorted, at times it can be helpful to sort a DataFrame so that the rows are ordered in a particular manner.

To do this, `sort_values` can be very helpful. For example, to sort by the values in 'col', you would use the following:

```python
df.sort_values(by=['col'])
```

To arrange (sort) the rows by more than one column, the syntax changes slightly:

```python
df.sort_values(by=['col1', 'col2'])
```

### Q9a: `sort_values` (0.25 points)

Suppose we want to know which individuals are oldest in our dataset.

**Sort the rows in data** such that the individual with the highest age is displayed first and the lowest age last.

Store the sorted DataFrame back in `data`.

Reset the index so that the oldest individual becomes index 0.

<!-- grade_id: cell-df2dae19d7386525, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9330a97562f3a48c, points: 0.25, grade: true, solution: false -->
```python
# Do not edit this cell. This cell is for testing purposes
```

## End of Assignment #1

When you think you're done, verify that everything is working correctly by doing the following:
1. Save this notebook
2. Choose the pull down menu item Kernel -> Restart and Run All Cells
3. Check that you see the print statement in the cell below. If you see that happy sunglasses face emoji, it means that you have passed all the publicly visible tests. Yay!!

```python
print("\N{SMILING FACE WITH SUNGLASSES}")
```